# CPSC 330 - Applied Machine Learning 

## Homework 3: Preprocessing 
### Associated lectures: [Lectures 4, 5, 6](https://github.com/UBC-CS/cpsc330-2023s/tree/main/lectures) 

**Due date: Monday, Feb 5, 11:59pm. See the [HW schedule](https://github.com/UBC-CS/cpsc330-2023W2#deliverable-due-dates-tentative)**. 


## Imports 

```python
from hashlib import sha1
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
from sklearn.compose import ColumnTransformer, make_column_transformer
from sklearn.dummy import DummyClassifier
from sklearn.impute import SimpleImputer
from sklearn.model_selection import cross_val_score, cross_validate, train_test_split
from sklearn.neighbors import KNeighborsClassifier
from sklearn.pipeline import Pipeline, make_pipeline
from sklearn.preprocessing import OneHotEncoder, StandardScaler, OrdinalEncoder
from sklearn.svm import SVC
from sklearn.tree import DecisionTreeClassifier
```

## Instructions
<hr>

_Points: 6_

Follow the [homework submission instructions](https://github.com/UBC-CS/cpsc330-2023W2/blob/master/docs/homework_instructions.md), and the summary at the end of this HW.

**You <u>may</u> work with a partner on this homework and submit your assignment as a group.** Below are some instructions on working as a group.  
- You can also work on your own if you prefer.
- The maximum group size is 2.
- Use group work as an opportunity to collaborate and learn new things from each other. 
- Be respectful to each other and make sure you understand all the concepts in the assignment well. 
- It's your responsibility to make sure that the assignment is submitted by one of the group members before the deadline. 

_Note: The assignments will get gradually more open-ended as we progress through the course. In many cases, there won't be a single correct solution. Sometimes you will have to make your own choices and your own decisions (for example, on what parameter values to use when they are not explicitly provided in the instructions). Use your own judgment in such cases and justify your choices, if necessary._

## Introduction <a name="in"></a>
<hr>

A crucial step when using machine learning algorithms on real-world datasets is preprocessing. This homework will give you some practice of data preprocessing and building a supervised machine learning pipeline on a real-world dataset. 

<br><br>

## Exercise 1: Introducing the dataset
<hr>

In this lab, you will be working on [the adult census dataset](https://www.kaggle.com/uciml/adult-census-income#). Download the CSV and save it as `adult.csv` under the data folder in this homework folder. 

This is a classification dataset and the classification task is to predict whether income exceeds 50K per year or not based on the census data. You can find more information on the dataset and features [here](http://archive.ics.uci.edu/ml/datasets/Adult).

The starter code below loads the data CSV (assuming that it is saved as `adult.csv` under the data folder). 

_Note that many popular datasets have sex as a feature where the possible values are male and female. This representation reflects how the data were collected and is not meant to imply that, for example, gender is binary._

```python
census_df = pd.read_csv("data/adult.csv")
census_df.shape
```

Output:
```
(32561, 15)
```

<br><br>

### 1.1 Data splitting 

_Points: 2_

In order to avoid violation of the golden rule, the first step before we do anything is splitting the data. 

**Your tasks:**

1. Split the data into `train_df` (40%) and `test_df` (60%) with `random_state = 123`. Keep the target column (`income`) in the splits so that we can use it in the exploratory data analysis.  

_Usually having more data for training is a good idea. But here I'm using 40%/60% split because running cross-validation with this dataset can take a while on a modest laptop. A smaller training data means it won't take too long to train the model on your laptop. A side advantage of this would be that with a bigger test split, we'll have a more reliable estimate of the model performance!_

<div class="alert alert-warning">

Solution_1.1
    
</div>

```python
train_df, test_df = train_test_split(census_df, train_size=0.4, test_size=0.6, random_state=123)

```

<br><br>

## Exercise 2: Exploratory data analysis (EDA) <a name="2"></a> 
<hr>

Let's examine our `train_df`. 

```python
train_df.sort_index()
```

Output:
```
age workclass  fnlwgt     education  education.num      marital.status  \
0       90         ?   77053       HS-grad              9             Widowed   
1       82   Private  132870       HS-grad              9             Widowed   
4       41   Private  264663  Some-college             10           Separated   
5       34   Private  216864       HS-grad              9            Divorced   
6       38   Private  150601          10th              6           Separated   
...    ...       ...     ...           ...            ...                 ...   
32554   32   Private  116138       Masters             14       Never-married   
32556   22   Private  310152  Some-college             10       Never-married   
32557   27   Private  257302    Assoc-acdm             12  Married-civ-spouse   
32559   58   Private  151910       HS-grad              9             Widowed   
32560   22   Private  201490       HS-grad              9       Never-married   

            occupation   relations
```

We see some missing values represented with a "?". Probably these were the questions not answered by some people during the census.  Usually `.describe()` or `.info()` methods would give you information on missing values. But here, they won't pick "?" as missing values because they are encoded as strings instead of an actual NaN in Python. So let's replace them with `np.nan` before we carry out EDA. If you do not do it, you'll encounter an error later on when you try to pass this data to a classifier. 

```python
train_df = train_df.replace("?", np.nan)
test_df = test_df.replace("?", np.nan)
train_df.shape
```

Output:
```
(13024, 15)
```

```python
train_df.sort_index()
```

Output:
```
age workclass  fnlwgt     education  education.num      marital.status  \
0       90       NaN   77053       HS-grad              9             Widowed   
1       82   Private  132870       HS-grad              9             Widowed   
4       41   Private  264663  Some-college             10           Separated   
5       34   Private  216864       HS-grad              9            Divorced   
6       38   Private  150601          10th              6           Separated   
...    ...       ...     ...           ...            ...                 ...   
32554   32   Private  116138       Masters             14       Never-married   
32556   22   Private  310152  Some-college             10       Never-married   
32557   27   Private  257302    Assoc-acdm             12  Married-civ-spouse   
32559   58   Private  151910       HS-grad              9             Widowed   
32560   22   Private  201490       HS-grad              9       Never-married   

            occupation   relations
```

The "?" symbols are now replaced with NaN values. 

<br><br>

### 2.1 `describe()` method

_Points: 5_

**Your tasks:**

1. Examine the output of `train_df.describe()` with `include='all'` argument and store it in a variable called `census_summary`.
2. What is the highest hours per week someone reported? Store it in a variable called `max_hours_per_week`.
3. What is the most frequently occurring occupation in this dataset? Store it in a variable called `most_freq_occupation`.
4. Store the column names of the columns with missing values as a list in a variable called `missing_vals_cols`. 
5. Store the column names of all numeric-looking columns as a list in a variable called `numeric_cols`.  

<div class="alert alert-warning">

Solution_2.1
    
</div>

```python
census_summary = train_df.describe(include='all')
census_summary
```

Output:
```
age workclass        fnlwgt education  education.num  \
count   13024.000000     12284  1.302400e+04     13024   13024.000000   
unique           NaN         8           NaN        16            NaN   
top              NaN   Private           NaN   HS-grad            NaN   
freq             NaN      9123           NaN      4153            NaN   
mean       38.546913       NaN  1.884759e+05       NaN      10.057432   
std        13.610225       NaN  1.040895e+05       NaN       2.553084   
min        17.000000       NaN  1.376900e+04       NaN       1.000000   
25%        28.000000       NaN  1.170962e+05       NaN       9.000000   
50%        37.000000       NaN  1.778990e+05       NaN      10.000000   
75%        48.000000       NaN  2.365655e+05       NaN      12.000000   
max        90.000000       NaN  1.184622e+06       NaN      16.000000   

            marital.status      occupation relationship   race    sex  \
count                13024           12281        13024  13024  130
```

```python
max_hours_per_week = train_df['hours.per.week'].max()
max_hours_per_week
```

Output:
```
99
```

```python
most_freq_occupation = train_df['occupation'].mode()
most_freq_occupation
```

Output:
```
0    Prof-specialty
Name: occupation, dtype: object
```

```python
col_has_null = train_df.isnull().any()
missing_vals_cols = col_has_null[col_has_null==True].index.tolist()
missing_vals_cols
numeric_cols = train_df.select_dtypes(include='number').columns.tolist()
numeric_cols
```

Output:
```
['age',
 'fnlwgt',
 'education.num',
 'capital.gain',
 'capital.loss',
 'hours.per.week']
```

<br><br>

<!-- BEGIN QUESTION -->

### 2.2 Visualizing features

_Points: 5_

**Your tasks:**

1. Visualize the histograms of numeric features.  
2. From the visualizations, which features seem relevant for the given prediction task?

<div class="alert alert-warning">

Solution_2.2
    
</div>

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 2.3 Identify transformations to apply

_Points: 13_

Before passing this data to a machine learning model, we need to apply some transformations on different features. Below we are providing possible transformations which can be applied on each column in `census_df`.  

**Your tasks:**
1. Write your justification or explanation for each row in the explanation column. An example explanation is given for the age column. 

> Note: This question is a bit open-ended. If you do not agree with the provided transformation, feel free to argue your case in the explanation. That said, in this assignment, go with the transformations provided below. 

<div class="alert alert-warning">

Solution_2.3
    
</div>

| Feature | Transformation | Explanation
| --- | ----------- | ----- |
| age | scaling |  A numeric feature with no missing values. Good idea to apply scaling, as the range of values (17 to 90) is quite different compared to other numeric features.|
| workclass | imputation, one-hot encoding | |
| fnlwgt | drop |  |
| education | ordinal encoding | |
| education.num | drop | |
| marital.status | one-hot encoding  | |
| occupation | imputation, one-hot encoding  | |
| relationship | one-hot encoding  | |
| race | drop  |  |
| sex | one-hot encoding with "binary=True" | |
| capital.gain | scaling |  | 
| capital.loss | scaling |  |
| hours.per.week | scaling | |
| native.country | imputation, one-hot encoding | | 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 2.4 Identify feature types 

_Points: 5_

**Your tasks:**
1. Based on the types of transformations you want to apply on the features, identify different feature types and store them in the variables below as lists.  

<div class="alert alert-warning">
    
Solution_2.4
    
</div>

```python
# Fill in the lists below.
numeric_features = []
categorical_features = []
ordinal_features = []
binary_features = []
drop_features = []
target = "income"

...
```

<br><br><br><br>

## Exercise 3: Baseline models 

### 3.1 Separating feature vectors and targets  

_Points: 2_

**Your tasks:**

1. Create `X_train`, `y_train`, `X_test`, `y_test` from `train_df` and `test_df`. 

<div class="alert alert-warning">
    
Solution_3.1
    
</div>

```python
X_train = None
y_train = None
X_test = None
y_test = None

...
```

<br><br>

### 3.2 Dummy classifier

_Points: 2_

**Your tasks:**

1. Carry out 5-fold cross-validation using `scikit-learn`'s `cross_validate` function with `return_train_scores=True` and store the results as a dataframe named `dummy_df` where each row corresponds to the results from a cross-validation fold. 

<div class="alert alert-warning">
    
Solution_3.2
    
</div>

```python
dummy_df = None 

...
```

<br><br>

<!-- BEGIN QUESTION -->

### 3.3 Discussion

_Points: 2_

**Your tasks:**

1. Hopefully, you were able to run cross-validation with dummy classifier successfully in the question above. At this point, if you train [`sklearn`'s `SVC`](https://scikit-learn.org/stable/modules/generated/sklearn.svm.SVC.html) model on `X_train` and `y_train` would it work? Why or why not? 

<div class="alert alert-warning">
    
Solution_3.3
    
</div>

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 4: Preprocessing
<hr>

In this dataset, we have different types of features: numeric features, an ordinal feature, categorical features, and a binary feature. We want to apply different transformations on different columns and therefore we need a column transformer. In this exercise, first, we'll define different transformations on different types of features and then will create a `scikit-learn`'s `ColumnTransformer`. For example, the code below creates a `numeric_transformer` for numeric features. 

```python
numeric_transformer = StandardScaler()
```

In the exercises below, you'll create transformers for other types of features. 

<br><br>

### 4.1 Preprocessing ordinal features

_Points: 5_

**Your tasks:**

1. Create a transformer called `ordinal_transformer` for our ordinal features. 

<div class="alert alert-warning">
    
Solution_4.1
    
</div>

```python
ordinal_transformer = None

...
```

```python
...
```

<br><br>

### 4.2 Preprocessing binary features

_Points: 2_

**Your tasks:**

1. Create a transformer called `binary_transformer` for our binary features.

> _Note that many popular datasets have sex as a feature where the possible values are male and female. This representation reflects how the data were collected and is not meant to imply that, for example, gender is binary._

<div class="alert alert-warning">
    
Solution_4.2
    
</div>

```python
binary_transformer = None
...
```

<br><br>

### 4.3 Preprocessing categorical features

_Points: 4_

In Exercise 2.3, we saw that there are 3 categorical features with missing values. So first we need to impute the missing values and then encode these features with one-hot encoding. For the purpose of this assignment, let's just have imputation as the first step for all categorical features even when they do not have missing values. This should be OK because if a feature doesn't have any missing value,  imputation won't be applied. 

If we want to apply more than one transformation on a set of features, we need to create a [`scikit-learn` `Pipeline`](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.Pipeline.html). For example, for categorical features we can create a `scikit-learn` `Pipeline` with first step as imputation and the second step as one-hot encoding. 

**Your tasks:**

1. Create a `sklearn` `Pipeline` using [`make_pipeline`](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.make_pipeline.html) called `categorical_transformer` for our categorical features with two steps: `SimpleImputer` for imputation with `strategy="constant"` and `fill_value="missing"` and `OneHotEncoder` with `handle_unknown="ignore"` and `sparse=False` for one-hot encoding. 

<div class="alert alert-warning">
    
Solution_4.3
    
</div>

```python
categorical_transformer = None

...
```

<br><br>

### 4.4 Creating a column transformer. 

_Points: 7_

**Your tasks:**
1. Create a `sklearn` `ColumnTransformer` named `preprocessor` using [`make_column_transformer`](https://scikit-learn.org/stable/modules/generated/sklearn.compose.make_column_transformer.html) with the transformers defined in the previous exercises. Use the sequence below in the column transformer and add a "drop" step for the `drop_features` in the end.  
    - `numeric_transformer`
    - `ordinal_transformer`
    - `binary_transformer`
    - `categorical_transformer`
2. Transform the data by calling `fit_transform` on the training set and save it as a dataframe in a variable called `transformed_df`. How many new columns have been created in the preprocessed data in comparison to the original `X_train`? Store the difference between the number of columns in `transformed_df` and `X_train` in a variable called `n_new_cols`. 

<div class="alert alert-warning">
    
Solution_4.4
    
</div>

```python
preprocessor = None

...
```

```python
transformed_df = None
n_new_cols = None

...
```

<br><br><br><br>

## Exercise 5: Building models 

Now that we have preprocessed features, we are ready to build models. Below, I'm providing the function we used in class which returns mean cross-validation score along with standard deviation for a given model. Use it to keep track of your results. 

```python
results_dict = {}  # dictionary to store all the results
```

```python
def mean_std_cross_val_scores(model, X_train, y_train, **kwargs):
    """
    Returns mean and std of cross validation

    Parameters
    ----------
    model :
        scikit-learn model
    X_train : numpy array or pandas DataFrame
        X in the training data
    y_train :
        y in the training data

    Returns
    ----------
        pandas Series with mean scores from cross_validation
    """

    scores = cross_validate(model, X_train, y_train, **kwargs)

    mean_scores = pd.DataFrame(scores).mean()
    std_scores = pd.DataFrame(scores).std()
    out_col = []

    for i in range(len(mean_scores)):
        out_col.append((f"%0.3f (+/- %0.3f)" % (mean_scores[i], std_scores[i])))

    return pd.Series(data=out_col, index=mean_scores.index)
```

Below I'm showing an example where I call `mean_std_cross_val_scores` with `DummyClassifier`. The function calls `cross_validate` with the passed arguments and returns a series with mean cross-validation results and std of cross-validation. When you train new models, you can just add the results of these models in `results_dict`, which can be easily converted to a dataframe so that you can have a table with all your results. 

```python
# Baseline model

from sklearn.dummy import DummyClassifier

dummy = DummyClassifier(random_state = 42)
pipe = make_pipeline(preprocessor, dummy)
results_dict["dummy"] = mean_std_cross_val_scores(
    pipe, X_train, y_train, cv=5, return_train_score=True
)
results_df = pd.DataFrame(results_dict).T
results_df
```

<br><br>

<!-- BEGIN QUESTION -->

### 5.1 Trying different classifiers

_Points: 10_

**Your tasks:**

1. For each of the models in the starter code below: 
    - Define a pipeline with two steps: `preprocessor` from 4.4 and the model as your classifier. 
    - Carry out 5-fold cross-validation with the pipeline and get the mean cross-validation scores with std by calling the `mean_std_cross_val_scores` function above. 
    - Store the results in a dataframe called `income_pred_results_df` with the model names in the `models` dictionary below as the index and each row representing results returned by `mean_std_cross_val_scores` function above. In other words, `income_pred_results_df` should look similar to the `results_df` dataframe above with more rows for the models below. 
    
> This might take a while to run. Be patient! 

```python
models = {
    "decision tree": DecisionTreeClassifier(random_state=42),
    "kNN": KNeighborsClassifier(),
    "RBF SVM": SVC(random_state=42),
}
```

<div class="alert alert-warning">
    
Solution_5.1
    
</div>

```python
income_pred_results_df = None 

...
```

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 5.2 Discussion 

_Points: 5_

**Your tasks:**

1. Examine the train and validation accuracies and `fit` and `score` times for all the models in the results above. How do the validation accuracies compare to the `DummyClassifier` model? Which model has the best validation accuracy? Which model is the fastest one? Which model is overfitting the most and the least?  


<div class="alert alert-warning">
    
Solution_5.2
    
</div>

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 5.3 Hyperparameter optimization

_Points: 10_

In this exercise, you'll carry out hyperparameter optimization for the hyperparameter `C` of SVC RBF classifier. In practice, you'll carry out hyperparameter optimization for all different hyperparameters of the most promising classifiers. For the purpose of this assignment, we'll only do it for the `SVC` classifier with one hyperparameter, namely `C`. 

**Your tasks:**

1. For each `C` value in the `param_grid` below: 
    - Create a pipeline object with two steps: preprocessor from 4.4 and `SVC` classifier with the `C` value.
    - Carry out 5-fold cross validation with the pipeline.  
    - Store the results in `results_dict` and display results as a pandas DataFrame. 
2. Which hyperparameter value seems to be performing the best? Store it in a variable called `best_C`. (Since this question is not autograded, please store the value directly as a number, something like `best_C = 0.001`, if `C = 0.001` is giving you the best CV score.) Is it different than the default value for the hyperparameter used by `scikit-learn`? 

> Note: Running this will take a while. Please be patient. 

```python
param_grid = {"C": np.logspace(-1, 2, 4)}
param_grid
```

<div class="alert alert-warning">
    
Solution_5.3
    
</div>

```python
...
```

```python
...
```

_Type your answer here, replacing this text._

```python
best_C = None

...
```

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 6: Evaluating on the test set 
<hr>

Now that we have a best performing model, it's time to assess our model on the set aside test set. In this exercise, you'll examine whether the results you obtained using cross-validation on the train set are consistent with the results on the test set. 

### 6.1 Scoring on the unseen test set 

_Points: 4_

**Your tasks:**

1. Create a pipeline named `final_pipeline` with the preprocessor from 4.4 as the first step and the best performing SVC model from 5.4 as the second step. 
2. Train the pipeline on the entire training set `X_train` and `y_train`. 
3. Score the pipeline on `X_test` and `y_test` and store the score in a variable called `test_score`.  

<div class="alert alert-warning">
    
Solution_6.1
    
</div>

```python
final_pipeline = None
test_score = None

...
```

_Type your answer here, replacing this text._

<br><br><br><br>

## Exercise 7: Short answer questions

<!-- BEGIN QUESTION -->

### 7.1 kNNs toy example

_Points: 3_

Suppose you want to get predictions using the $k$-nearest neighbour algorithm on the toy dataset below. 

$$ X = \begin{bmatrix}5 & 2\\4 & -2\\  2 & 2\\ 10 & 10\\ 9 & -1\\ 9& 9\end{bmatrix}, \quad y = \begin{bmatrix}0\\0\\1\\1\\1\\2\end{bmatrix}.$$

Suppose you are given a new data point $x=\begin{bmatrix} 0\\0\end{bmatrix}$.  

1. What would the model predict when $k=1$?
2. What would the model predict when $k=3$?
3. What would the model predict when $k=3$ if we were doing regression rather than classification? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 7.2 Preprocessing

_Points: 8_

1. What is the problem with calling `fit_transform` on your test data with `StandardScaler`?
2. Why is it important to follow the Golden Rule? If you violate it, will that give you a worse classifier?
3. What are two advantages of using sklearn Pipelines? 
4. When is it appropriate to use sklearn `ColumnTransformer`? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

## Submission instructions 

**PLEASE READ:** When you are ready to submit your assignment do the following:

1. Run all cells in your notebook to make sure there are no errors by doing `Kernel -> Restart Kernel and Clear All Outputs` and then `Run -> Run All Cells`. 
2. Notebooks with cell execution numbers out of order or not starting from “1” will have marks deducted. Notebooks without the output displayed may not be graded at all (because we need to see the output in order to grade your work).
3. Upload the assignment using [PrairieLearn](https://ca.prairielearn.com/pl/course_instance/6697).

Congratulations on finishing the homework! This was a tricky one but I hope you are feeling good after working on it. You are now ready to build a simple supervised machine learning pipeline on real-world datasets! Well done :clap:! 

![](img/eva-well-done.png)



```python

```

Tests for q1.1 (points: 2)
```
>>> assert not train_df is None and not test_df is None, "Please use the provided variables."
>>> assert train_df.shape == (13024, 15), "The dimensions of the training set are incorrect"
>>> assert test_df.shape == (19537, 15), "The dimensions of the test set are incorrect"
>>> assert train_df.loc[12846][['age', 'education', 'occupation', 'capital.loss']].tolist() == [49, 'Some-college', 'Craft-repair', 0], "Are you using the provided random state?"
>>> assert not 20713 in train_df.index, 'Are you using the provided random state?'
```


Tests for q2.1 (points: [1, 1, 1, 1, 1])
```
>>> # Task 1
>>> assert isinstance(census_summary, pd.DataFrame), "census_summary dataftame is not created"
>>> assert census_summary.shape == (11, 15), "census_summary shape is incorrect. Probably you are not including all columns"
>>> assert census_summary.loc['min']['age'] == 17.0, "census_summary dataframe is incorrect"
>>> assert census_summary.loc['top']['occupation'] == "Prof-specialty", "census_summary dataframe is incorrect"
```
```
>>> # Task 2
>>> assert (sha1(str(max_hours_per_week).encode('utf8')).hexdigest() == "3359de52c8ae993fe0f8fe9c5168a0065bd3c7a4"), "max_hours_per_week are incorrect"
```
```
>>> # Task 3
>>> assert (sha1(str(most_freq_occupation).encode('utf8')).hexdigest() == "97165f50eddb0d28a382b0366274e2fe38505644"), "most_freq_occupation is incorrect"
```
```
>>> # Task 4
>>> assert (sha1(str(missing_vals_cols).encode('utf8')).hexdigest() == "6bc5e13d4d66b306e52701ee9a1e5e21bf19aeb0"), "Please use the exact column/feature name. Also, make sure the lists are sorted."
```
```
>>> # Task 5
>>> assert (sha1(str(numeric_cols).encode('utf8')).hexdigest() == "615afaf5011128d641ab8a73289d57bd01a3ec37"), "Please use the exact column/feature name. Also, make sure the lists are sorted."
```


Tests for q2.4 (points: 5)
```
>>> assert (sha1(str(numeric_features).encode('utf8')).hexdigest() == "71401cf60034fd69eee7398866359f612adf3e15"), "numeric_features list is not correct"
>>> assert (sha1(str(categorical_features).encode('utf8')).hexdigest() == "af1a4022c0362405678be5c3a6735578a8c0069f"), "categorical_features list is not correct"
>>> assert (sha1(str(ordinal_features).encode('utf8')).hexdigest() == "95b86602c44211f3ad662bb58b8e53d024106d05"), "ordinal_features list is not correct"
>>> assert (sha1(str(binary_features).encode('utf8')).hexdigest() == "d4b7aa4c56ac2f98e6ac9cec7768484b415b7337"), "binary_features list is not correct"
>>> assert (sha1(str(drop_features).encode('utf8')).hexdigest() == "62aab57d42c54be3dfd3c55020e5a167ca1a84c3"), "drop_features list is not correct"
>>> assert (sha1(str(target).encode('utf8')).hexdigest() == "0f613350b66e64d92ef21bc4dcdbf8996cb4edf0"), "target variable is not set correctly"
```


Tests for q3.1 (points: 2)
```
>>> assert not X_train is None, "Your answer does not exist. Have you passed in the correct variable?"
>>> assert not y_train is None, "Your answer does not exist. Have you passed in the correct variable?"
>>> assert not X_test is None, "Your answer does not exist. Have you passed in the correct variable?"
>>> assert not y_test is None, "Your answer does not exist. Have you passed in the correct variable?"
>>> assert X_train.shape == (13024, 14), "The dimensions of X_train are incorrect"
>>> assert y_train.shape == (13024, ), "The dimensions of y_train are incorrect. Are you splitting correctly"
>>> assert X_test.shape == (19537,14), "The dimensions of X_test are incorrect. Are you splitting correctly? Are you using single brackets?"
>>> assert y_test.shape == (19537,), "The dimensions of y_test are incorrect. Are you splitting correctly? Are you using single brackets?"
>>> assert 'income' not in list(X_train.columns), "Make sure the target variable is not part of your X dataset."
```


Tests for q3.2 (points: 2)
```
>>> assert not dummy_df is None, "Have you used the correct variable to store the results?"
>>> assert sorted(list(dummy_df.columns)) == ['fit_time','score_time','test_score','train_score'], "Your solution contains incorrect columns."
>>> assert dummy_df.shape == (5,4), "Are you carrying out 5-fold cross-validation and are you passing return_train_score=True?"
>>> assert sha1(str(round(dummy_df['test_score'].mean(),3)).encode('utf8')).hexdigest() == "e04884a1f90ee71d58aa5550207e305cd0e18392", "The test scores seem wrong. Are you calling cross_validate correctly?"
>>> assert sha1(str(round(dummy_df['train_score'].mean(),3)).encode('utf8')).hexdigest() == "e04884a1f90ee71d58aa5550207e305cd0e18392", "The train scores seem wrong. Are you calling cross_validate correctly?"
```


Tests for q4.1 (points: 5)
```
>>> assert not ordinal_transformer is None, "Are you using the correct variable name?"
>>> assert type(ordinal_transformer.get_params()['categories'][0]) is list, "Are you passing education levels as a list of lists?"
>>> assert ordinal_transformer.get_params()['dtype'] == int, "Please set the dtype to int"
>>> assert (sha1(str(ordinal_transformer.get_params()['categories'][0]).encode('utf8')).hexdigest() == "893a03d114b2af09b53247866c6eea54ebfd090f") or (sha1(str(ordinal_transformer.get_params()['categories'][0]).encode('utf8')).hexdigest() == "81059b8bebc9ddb03d61bf07cfd9b9b6b0da288e"), "Make sure you are passing categories sorted on levels of education. (Ascending or descending shouldn't matter.)"
```


Tests for q4.2 (points: 2)
```
>>> assert not binary_transformer is None, "Are you using the correct variable name?"
>>> assert binary_transformer.get_params()['drop'] == 'if_binary', "Are you passing `drop=if_binary`?"
>>> assert binary_transformer.get_params()['dtype'] == int, "Please set the dtype to int"
```


Tests for q4.3 (points: 4)
```
>>> assert not categorical_transformer is None, "Are you using the correct variable name?"
>>> assert type(categorical_transformer) is Pipeline, "Are you creating a scikit-learn Pipeline?"
>>> assert len(categorical_transformer.get_params()['steps']) == 2, "Are you creating a pipeline with two steps?"
>>> assert categorical_transformer.get_params()['simpleimputer__strategy'] == 'constant', "Are you passing strategy=constant in the SimpleImputer?"
>>> assert categorical_transformer.get_params()['simpleimputer__fill_value'] == 'missing', "Are you passing fill_value='missing' in the SimpleImputer?"
>>> assert categorical_transformer.get_params()['onehotencoder__handle_unknown'] == 'ignore', "Are you passing handle_unknown = 'ignore' argument to your OHE?"
>>> assert categorical_transformer.get_params()['onehotencoder__sparse'] == False, "Are you creating a sparse matrix for OHE?"
```


Tests for q4.4 (points: [5, 1, 1])
```
>>> # task 1
>>> assert not preprocessor is None, "Are you using the correct variable name?"
>>> assert len(preprocessor.get_params()['transformers']) in range(4,6,1), "Have you included all the transformers?"
>>> assert 'onehotencoder' in preprocessor.get_params().keys(), 'Either the categorical_transformer or binary_transformer is not included.'
>>> assert 'standardscaler' in preprocessor.get_params().keys(), 'numeric_transformer is not included.'
>>> assert 'ordinalencoder' in preprocessor.get_params().keys(), 'ordinal_transformer is not included.'
>>> assert 'drop' in preprocessor.get_params().keys(), 'drop features step is not included.'
```
```
>>> # task 2
>>> assert not transformed_df is None, "Are you using the correct variable name?"
>>> assert sha1(str(transformed_df.shape).encode('utf8')).hexdigest() == 'a0521f0cdbcd77cd213e7d1a3cfc13c1c7c92a6e', "The shape of the transformed data is incorrect."
```
```
>>> assert sha1(str(n_new_cols).encode('utf8')).hexdigest() == 'b7103ca278a75cad8f7d065acda0c2e80da0b7dc', "The number of new columns (n_new_cols) is incorrect."
```


Tests for q6.1 (points: [2, 2])
```
>>> # task 1
>>> assert not final_pipeline is None, "Are you using the correct variable name?"
>>> assert not test_score is None, "Are you using the correct variable name?"
>>> assert len(final_pipeline.named_steps) == 2, "The final pipeline needs to have two steps: one for the preprocessor and one for SVC."
>>> assert final_pipeline.n_features_in_ == 14, "Make sure to pass the original X_train to fit"
>>> assert final_pipeline.named_steps['svc'].get_params()['C'] == best_C, "Are you using the best C value from the previous exercise?"
```
```
>>> # task 2
>>> assert not test_score is None, "Are you using the correct variable name?"
>>> assert sha1(str(round(test_score,3)).encode('utf8')).hexdigest() == '1284d7d1d642ef8e51475cce5c1972ca6b8bd2b6', "The test score seems off"
```
